# HW2: Khám phá dữ liệu, tiền xử lý, phân tích đơn giản
**Vì đây là bài tập về Pandas nên yêu cầu là không được dùng vòng lặp**

(Cập nhật lần cuối: 25/07/2021)

Họ tên: Lê Thành Nam

MSSV: 18127158

---

## Cách làm bài và nộp bài (bạn cần đọc kỹ)

&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác cũng như tham khảo các nguồn trên mạng, nhưng sau cùng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. Khi tham khảo các nguồn trên mạng thì bạn cần ghi rõ nguồn trong bài làm. Bạn không được tham khảo bài làm của các bạn năm trước (vì nếu làm vậy thì bao giờ bạn mới có thể tự mình suy nghĩ để giải quyết vấn đề); sau khi kết thúc môn học, bạn cũng không được đưa bài làm cho các bạn khóa sau hoặc public bài làm trên Github (vì nếu làm vậy thì sẽ ảnh hưởng tới việc học của các bạn khóa sau). Nếu bạn có thể làm theo những gì mình nói thì điểm của bạn có thể sẽ không cao nhưng bạn sẽ có được những bước tiến thật sự. <font color=red>Trong trường hợp bạn vi phạm những điều mình nói ở trên thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart & Run All`, để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart & Run All` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - File `HW2.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. Đuôi của file nén phải là .zip (chứ không được .rar hay gì khác).

<font color=red>Bạn lưu ý tuân thủ chính xác qui định nộp bài ở trên.</font>

---

## Môi trường code

Ta thống nhất trong môn này: dùng phiên bản các package như trong file "min_ds-env.yml" (file này đã được cập nhật một số lần, file mới nhất ở đầu có dòng: "Last update: 24/06/2021"). Cách tạo ra môi trường code từ file "min_ds-env.yml" đã được nói ở file "02_BeforeClass-Notebook_Python.pdf".

Check môi trường code:

```python
import sys
sys.executable
```

Output:
```
'D:\\Works\\Anaconda\\envs\\min_ds-env\\python.exe'
```

Nếu không có vấn đề gì thì file chạy python sẽ là file của môi trường code "min_ds-env".

---

## Import

```python
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd 
import operator 
```

---

## Thu thập dữ liệu

Dữ liệu được sử dụng trong bài tập này là dữ liệu khảo sát các lập trình viên của trang StackOverflow. Mình download dữ liệu [ở đây](https://drive.google.com/file/d/1dfGerWeWkcyQ9GX9x20rdSGj7WtEpzBB/view) và có bỏ đi một số cột để đơn giản hóa. 

Theo mô tả trong file "README_2020.txt" của StackOverflow:
>The enclosed data set is the full, cleaned results of the 2020 Stack Overflow Developer Survey. Free response submissions and personally identifying information have been removed from the results to protect the privacy of respondents. There are three files besides this README:
>
>1. survey_results_public.csv - CSV file with main survey results, one respondent per row and one column per answer
>2. survey_results_schema.csv - CSV file with survey schema, i.e., the questions that correspond to each column name
>3. so_survey_2020.pdf - PDF file of survey instrument
>
>The survey was fielded from February 5 to February 28, 2020. The median time spent on the survey for qualified responses was 16.6 minutes.
>
>Respondents were recruited primarily through channels owned by Stack Overflow. The top 5 sources of respondents were onsite messaging, blog posts, email lists, Meta posts, banner ads, and social media posts. Since respondents were recruited in this way, highly engaged users on Stack Overflow were more likely to notice the links for the survey and click to begin it.

File "survey_results_public-short.csv" mà mình đính kèm là phiên bản đơn giản hóa của file "survey_results_public.csv" (từ 61 cột, mình bỏ xuống còn 29 cột). Đây là file dữ liệu chính mà bạn sẽ làm trong bài tập này. Ngoài ra, mình còn đính kèm 2 file phụ: (1) file "survey_results_schema-short.csv" là file cho biết ý nghĩa của các cột, và (2) file "so_survey_2020.pdf" là file khảo sát gốc của StackOverflow.

Để ý: 
- Dữ liệu này không đại diện được cho cộng đồng lập trình viên trên toàn thế giới, mà chỉ giới hạn trong tập những lập trình viên thực hiện khảo sát của StackOverflow. Những câu trả lời có được thông qua tập dữ liệu này cũng sẽ bị giới hạn trong phạm vi đó.
- Dữ liệu có đúng không? Về cơ bản là ta không biết được. Ở đây, mục đích chính là học qui trình Khoa Học Dữ Liệu và các câu lệnh của Pandas nên ta sẽ **giả định** phần lớn dữ liệu là đúng và tiếp tục làm.

Cũng theo file "README_2020.txt", dữ liệu này được StackOverflow public với license như sau:
>This database - The Public 2020 Stack Overflow Developer Survey Results - is made available under the Open Database License (ODbL): http://opendatacommons.org/licenses/odbl/1.0/. Any rights in individual contents of the database are licensed under the Database Contents License: http://opendatacommons.org/licenses/dbcl/1.0/
>
>TLDR: You are free to share, adapt, and create derivative works from The Public 2020 Stack Overflow Developer Survey Results as long as you attribute Stack Overflow, keep the database open (if you redistribute it), and continue to share-alike any adapted database under the ODbl.


---

## Khám phá dữ liệu

### Đọc dữ liệu từ file (0.25đ)

Đầu tiên, bạn viết code để đọc dữ liệu từ file "survey_results_public-short.csv" và lưu kết quả vào DataFrame `survey_df`; ta thống nhất là sẽ để file dữ liệu này cùng cấp với file notebook và khi đọc file thì chỉ truyền vào tên của file. Ngoài ra, bạn cũng cần cho cột `Respondent` (id của người làm khảo sát) làm cột index của `survey_df`.

<!-- grade_id: cell-d46e97079096473c, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
file_path = "D:/Works/KHDL/survey_results_public-short.csv"
survey_df = pd.read_csv(file_path, index_col = 'Respondent')
```

```python
# TEST
survey_df.head()
```

Output:
```
MainBranch Hobbyist   Age  \
Respondent                                                                     
1                              I am a developer by profession      Yes   NaN   
2                              I am a developer by profession       No   NaN   
3                                 I code primarily as a hobby      Yes   NaN   
4                              I am a developer by profession      Yes  25.0   
5           I used to be a developer by profession, but no...      Yes  31.0   

            ConvertedComp             Country  \
Respondent                                      
1                     NaN             Germany   
2                     NaN      United Kingdom   
3                     NaN  Russian Federation   
4                     NaN             Albania   
5                     NaN       United States   

                                                      DevType  \
Respondent                                                      
1           Deve
```

### Dữ liệu có bao nhiêu dòng và bao nhiêu cột? (0.25đ)

Kế đến, bạn tính số dòng và số cột của DataFrame `survey_df` và lần lượt lưu vào biến `num_rows` và `num_cols`. 

<!-- grade_id: cell-8338f3610f6c5002, points: 0, grade: false, solution: true -->
```python
(num_rows, num_cols) = survey_df.shape 
```

<!-- grade_id: cell-00f00d1bba235517, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert num_rows == 64461
assert num_cols == 28
```

### Mỗi dòng có ý nghĩa gì? Có vấn đề các dòng có ý nghĩa khác nhau không?

Theo file "README_2020.txt" cũng như theo quan sát sơ bộ về dữ liệu, mỗi dòng trong DataFrame `survey_df` cho biết kết quả làm khảo sát của một người. Có vẻ không có vấn đề các dòng có ý nghĩa khác nhau. 

### Dữ liệu có các dòng bị lặp không? (0.25đ)

Kế đến, bạn tính số dòng có index (id của người làm khảo sát) bị lặp và lưu vào biến `num_duplicated_rows`. Trong nhóm các dòng có index giống nhau thì dòng đầu tiên không tính là bị lặp.

<!-- grade_id: cell-aadbfe12be205faa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
duplicated = survey_df.index.duplicated()
true_numpy = [i for i in duplicated if i == True]
num_duplicated_rows = len(true_numpy)
```

<!-- grade_id: cell-7c456d5495d4904d, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert num_duplicated_rows == 0
```

### Mỗi cột có ý nghĩa gì? (0.25đ)

Để xem ý nghĩa của mỗi cột thì:
- Trước tiên, bạn cần đọc file "survey_results_schema-short.csv" vào DataFrame `col_meaning_df`; bạn cũng cần cho cột "Column" làm cột index. 
- Sau đó, bạn chỉ cần hiển thị DataFrame `col_meaning_df` ra để xem (vụ này khó nên ở dưới mình đã làm cho bạn ở cell có dòng "# TEST" 😉). Tuy nhiên, bạn sẽ thấy ở cột "QuestionText": các chuỗi mô tả bị cắt do quá dài. Do đó, trước khi hiển thị DataFrame `col_meaning_df`, bạn cũng cần chỉnh sao đó để các chuỗi mô tả không bị cắt (vụ này bạn tự search Google, gợi ý: bạn sẽ dùng đến câu lệnh `pd.set_option`).   

<!-- grade_id: cell-5efb0517df207535, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
file_path = "D:/Works/KHDL/survey_results_schema-short.csv"
col_meaning_df = pd.read_csv(file_path, index_col = "Column")
```

```python
# TEST
col_meaning_df
```

Output:
```
QuestionText
Column                                                                         
Respondent                    Randomized respondent ID number (not in order ...
MainBranch                    Which of the following options best describes ...
Hobbyist                                                Do you code as a hobby?
Age                           What is your age (in years)? If you prefer not...
ConvertedComp                 Salary converted to annual USD salaries using ...
Country                                                      Where do you live?
DevType                       Which of the following describe you? Please se...
EdLevel                       Which of the following best describes the high...
Employment                    Which of the following best describes your cur...
Ethnicity                     Which of the following describe you, if any? P...
Gender                        Which of the following describe you, if any? P...
JobFactors                 
```

Trước khi đi tiếp, bạn nên đọc kết quả hiển thị ở trên và đảm bảo là bạn đã hiểu ý nghĩa của các cột. Để hiểu ý nghĩa của cột, có thể bạn sẽ cần xem thêm các giá trị của cột bên DataFrame `survey_df`. 

### Mỗi cột hiện đang có kiểu dữ liệu gì? Có cột nào có kiểu dữ liệu chưa phù hợp để có thể xử lý tiếp không? (0.25đ)

Kế đến, bạn tính kiểu dữ liệu (dtype) của mỗi cột trong DataFrame `survey_df` và lưu kết quả vào Series `dtypes` (Series này có index là tên cột). 

<!-- grade_id: cell-206c19691b1a6c05, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
dtypes = survey_df.dtypes
dtypes
```

Output:
```
MainBranch                       object
Hobbyist                         object
Age                             float64
ConvertedComp                   float64
Country                          object
DevType                          object
EdLevel                          object
Employment                       object
Ethnicity                        object
Gender                           object
JobFactors                       object
JobSat                           object
LanguageDesireNextYear           object
LanguageWorkedWith               object
MiscTechDesireNextYear           object
MiscTechWorkedWith               object
NEWCollabToolsDesireNextYear     object
NEWCollabToolsWorkedWith         object
NEWEdImpt                        object
NEWLearn                         object
NEWOvertime                      object
NEWStuck                         object
OpSys                            object
PlatformDesireNextYear           object
PlatformWorkedWith               object

```

<!-- grade_id: cell-3ebcf8f07733605b, points: 0.25, grade: true, solution: false -->
```python
# TEST
float_cols = set(dtypes[(dtypes==np.float32) | (dtypes==np.float64)].index)
assert float_cols == {'Age', 'ConvertedComp', 'WorkWeekHrs'}
object_cols = set(dtypes[dtypes == object].index)
assert len(object_cols) == 25
```

Như bạn có thể thấy, cột "YearsCode" và "YearsCodePro" nên có kiểu dữ liệu số, nhưng hiện giờ đang có kiểu dữ liệu object. Ta hãy thử xem thêm về các giá trị 2 cột này. 

```python
survey_df['YearsCode'].unique()
```

Output:
```
array(['36', '7', '4', '15', '6', '17', '8', '10', '35', '5', '37', '19',
       '9', '22', '30', '23', '20', '2', 'Less than 1 year', '3', '13',
       '25', '16', '43', '11', '38', '33', nan, '24', '21', '12', '40',
       '27', '50', '46', '14', '18', '28', '32', '44', '26', '42', '31',
       '34', '29', '1', '39', '41', '45', 'More than 50 years', '47',
       '49', '48'], dtype=object)
```

```python
survey_df['YearsCodePro'].unique()
```

Output:
```
array(['27', '4', nan, '8', '13', '2', '7', '20', '1', '23', '3', '12',
       '17', '18', '10', '14', '29', '6', '28', '9', '15', '11', '16',
       '25', 'Less than 1 year', '5', '21', '19', '35', '24', '32', '22',
       '30', '38', '26', '40', '33', '31', 'More than 50 years', '34',
       '36', '39', '37', '41', '45', '47', '42', '46', '50', '43', '44',
       '48', '49'], dtype=object)
```

Ta nên đưa 2 cột này về dạng số để có thể tiếp tục khám phá (tính min, median, max, ...).

---

## Tiền xử lý (0.5đ)

Bạn sẽ thực hiện tiền xử lý để chuyển 2 cột "YearsCode" và "YearsCodePro" về dạng số (float). Trong đó: "Less than 1 year" $\to$ 0, "More than 50 years" $\to$ 51. Sau khi chuyển thì `survey_df.dtypes` sẽ thay đổi.

<!-- grade_id: cell-cfc7d3af0fb5edec, points: 0, grade: false, solution: true -->
```python
# 
survey_df['YearsCode'] = survey_df.YearsCode.replace(["Less than 1 year", "More than 50 years"], ["0", "51"]).astype(np.float64)

survey_df['YearsCodePro'] = survey_df.YearsCodePro.replace(["Less than 1 year", "More than 50 years"], ["0", "51"]).astype(np.float64)

#survey_df.dtypes
survey_df.dtypes
```

Output:
```
MainBranch                       object
Hobbyist                         object
Age                             float64
ConvertedComp                   float64
Country                          object
DevType                          object
EdLevel                          object
Employment                       object
Ethnicity                        object
Gender                           object
JobFactors                       object
JobSat                           object
LanguageDesireNextYear           object
LanguageWorkedWith               object
MiscTechDesireNextYear           object
MiscTechWorkedWith               object
NEWCollabToolsDesireNextYear     object
NEWCollabToolsWorkedWith         object
NEWEdImpt                        object
NEWLearn                         object
NEWOvertime                      object
NEWStuck                         object
OpSys                            object
PlatformDesireNextYear           object
PlatformWorkedWith               object

```

<!-- grade_id: cell-4ee41f84604837a1, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert survey_df['YearsCode'].dtype in [np.float32, np.float64]
assert survey_df['YearsCodePro'].dtype in [np.float32, np.float64]
```

---

## Quay lại bước khám phá dữ liệu 

### Với mỗi cột có kiểu dữ liệu dạng số, các giá trị được phân bố như thế nào? (1đ)

(Trong đó: phần tính các mô tả của mỗi cột chiếm 0.5đ, phần tính số lượng giá trị không hợp lệ của mỗi cột chiếm 0.5đ)

Với các cột có kiểu dữ liệu số, bạn sẽ tính:
- Tỉ lệ % (từ 0 đến 100) các giá trị thiếu 
- Giá trị min
- Giá trị lower quartile (phân vị 25)
- Giá trị median (phân vị 50)
- Giá trị upper quartile (phân vị 75)
- Giá trị max

Bạn sẽ lưu kết quả vào DataFrame `nume_col_info_df`, trong đó: 
- Tên của các cột là tên của các cột số trong `survey_df`
- Tên của các dòng là: "missing_percentage", "min", "lower_quartile", "median", "upper_quartile", "max"  

Để dễ nhìn, tất cả các giá trị bạn đều làm tròn với 1 chữ số thập phân bằng phương thức `.round(1)`.

<!-- grade_id: cell-d1503f01712c873a, points: 0, grade: false, solution: true -->
```python
def get_numeric_col(survey_df): #Trả về list chứa tên các columns chứa giá trị số
    dtype = survey_df.dtypes
    result = []
    for i in range(len(dtype)):
        if dtype[i] == "float64" or dtype[i] =="float32" or dtype[i] == "int32" or dtype[i] == "int64":
            result.append(dtype.index[i])
    return result
    
def nume_col_info_df(survey_df, numeric_cols, cal_index): #Tính toán 5 giá trị trên cho toàn bộ columns có giá trị số
    result_df = pd.DataFrame(columns = numeric_cols, index = cal_index)
    for _index, row in result_df.iterrows():
        for _col in numeric_cols:
            if _index == "missing_percentage":
                row[_col] = (survey_df[_col].isna().sum() / survey_df[_col].size * 100).round(1)
            elif _index == "min":  
                row[_col] = survey_df[_col].min(skipna = True).round(1)
            elif _index == "lower_quartile":
                row[_col] = survey_df[_col].quantile([0.25]).values[0].round(1)
            elif _index == "median":
                row[_col] = survey_df[_col].quantile([0.5]).values[0].round(1)
            elif _index == "upper_quartile":
                row[_col] = survey_df[_col].quantile([0.75]).values[0].round(1)
            elif _index == "max":
                row[_col] = survey_df[_col].max(skipna = True).round(1)                
    return result_df

numeric_cols = get_numeric_col(survey_df)
cal_index = ["missing_percentage", "min", "lower_quartile", "median", "upper_quartile", "max"]

nume_col_info_df = nume_col_info_df(survey_df, numeric_cols, cal_index)
nume_col_info_df
```

Output:
```
Age ConvertedComp WorkWeekHrs YearsCode YearsCodePro
missing_percentage   29.5          46.1        36.2      10.5         28.1
min                   1.0           0.0         1.0       0.0          0.0
lower_quartile       24.0       24648.0        40.0       6.0          3.0
median               29.0       54049.0        40.0      10.0          6.0
upper_quartile       35.0       95000.0        44.0      17.0         12.0
max                 279.0     2000000.0       475.0      51.0         51.0
```

<!-- grade_id: cell-b19aff83344102f8, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert nume_col_info_df.shape == (6, 5)
data = nume_col_info_df.loc[['missing_percentage', 'min', 'lower_quartile', 'median', 'upper_quartile', 'max'],
                           ['Age', 'ConvertedComp', 'WorkWeekHrs', 'YearsCode', 'YearsCodePro']].values
correct_data = np.array([[     29.5,      46.1,      36.2,      10.5,      28.1],
                         [      1. ,       0. ,       1. ,       0. ,       0. ],
                         [     24. ,   24648. ,      40. ,       6. ,       3. ],
                         [     29. ,   54049. ,      40. ,      10. ,       6. ],
                         [     35. ,   95000. ,      44. ,      17. ,      12. ],
                         [    279. , 2000000. ,     475. ,      51. ,      51. ]])
assert np.array_equal(data, correct_data)
```

**Có giá trị không hợp lệ trong mỗi cột không? (không xét giá trị thiếu)**
- Cột "Age": bạn hãy tính số lượng giá trị không hợp lệ của cột "Age" (< giá trị tương ứng trong cột "YearsCode" HOẶC < giá trị tương ứng trong cột "YearsCodePro") và lưu kết quả vào biến `num_invalid_Age_vals`.
- Cột "WorkWeekHrs" (số giờ làm việc trung bình một tuần): ta thấy max là 475 giờ! Trong khi đó, 7 ngày * 24 giờ = 168 giờ! Bạn hãy tính số lượng giá trị không hợp lệ của cột "WorkWeekHrs" (> 24 * 7) và lưu kết quả vào biến `num_invalid_WorkWeekHrs`.
- Cột "YearsCode": bạn hãy tính số lượng giá trị không hợp lệ của cột "YearsCode" (< giá trị tương ứng trong cột "YearsCodePro" HOẶC > giá trị tương ứng trong cột "Age") và lưu kết quả vào biến `num_invalid_YearsCode`.
- Cột "YearsCodePro": bạn hãy tính số lượng giá trị không hợp lệ của cột "YearsCodePro" (> giá trị tương ứng trong cột "YearsCode" HOẶC > giá trị tương ứng trong cột "Age") và lưu kết quả vào biến `num_invalid_YearsCodePro`.

<!-- grade_id: cell-59a58bdc5f54bb25, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
num_invalid_Age_vals = survey_df[(survey_df['Age'] < survey_df['YearsCode']) | (survey_df['Age'] < survey_df['YearsCodePro'])].shape[0]
num_invalid_WorkWeekHrs_vals = survey_df[survey_df['WorkWeekHrs'] > 24 * 7].shape[0]
num_invalid_YearsCode_vals = survey_df[(survey_df['YearsCode'] < survey_df['YearsCodePro']) | (survey_df['YearsCode'] > survey_df['Age'])].shape[0]
num_invalid_YearsCodePro_vals = survey_df[(survey_df['YearsCodePro'] > survey_df['YearsCode']) | (survey_df['YearsCodePro'] > survey_df['Age'])].shape[0]

```

<!-- grade_id: cell-386769471c5b4b8f, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert num_invalid_WorkWeekHrs_vals == 62
assert num_invalid_Age_vals == 16
assert num_invalid_YearsCode_vals == 499
assert num_invalid_YearsCodePro_vals == 486
```

Do số lượng các giá trị không hợp lệ cũng khá ít nên ta có thể tiền xử lý bằng cách xóa các dòng chứa các giá trị không hợp lệ.

---

## Tiền xử lý (0.5đ)

Bạn sẽ thực hiện tiền xử lý để xóa đi các dòng của DataFrame `survey_df` mà chứa ít nhất là một giá trị không hợp lệ. Sau khi tiền xử lý thì `survey_df` sẽ thay đổi. 

<!-- grade_id: cell-6a6242ef6d2c8c65, points: 0, grade: false, solution: true -->
```python

#4 điều kiện lọc dataframe
invalid_Age = (survey_df['Age'] < survey_df['YearsCode']) | (survey_df['Age'] < survey_df['YearsCodePro'])
invalid_WorkWeekHrs = survey_df['WorkWeekHrs'] > 24 * 7
invalid_YearsCode = (survey_df['YearsCode'] < survey_df['YearsCodePro']) | (survey_df['YearsCode'] > survey_df['Age'])
invalid_YearsCodePro = (survey_df['YearsCodePro'] > survey_df['YearsCode']) | (survey_df['YearsCodePro'] > survey_df['Age'])
drop_index = survey_df[invalid_Age | invalid_WorkWeekHrs | invalid_YearsCode | invalid_YearsCodePro].index
survey_df.drop(drop_index, inplace = True)
len(survey_df)
```

Output:
```
63900
```

<!-- grade_id: cell-9d04f770468db3c8, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert len(survey_df) == 63900
```

---

## Quay lại bước khám phá dữ liệu

### Với mỗi cột có kiểu dữ liệu không phải dạng số, các giá trị được phân bố như thế nào? (1đ)

Với các cột có kiểu dữ liệu không phải số, bạn sẽ tính:
- Tỉ lệ % (từ 0 đến 100) các giá trị thiếu 
- Số lượng các giá trị (các giá trị ở đây là các giá trị khác nhau và ta không xét giá trị thiếu): với cột mà ứng với câu hỏi dạng multichoice (ví dụ, cột "DevType"), mỗi giá trị có thể chứa nhiều choice (các choice được phân tách bởi dấu chấm phẩy), và việc đếm trực tiếp các giá trị không có nhiều ý nghĩa lắm vì số lượng tổ hợp các choice là khá nhiều; một cách khác tốt hơn mà bạn sẽ làm là đếm số lượng các choice
- Tỉ lệ % (từ 0 đến 100) của mỗi giá trị được sort theo tỉ lệ % giảm dần (ta không xét giá trị thiếu, tỉ lệ là tỉ lệ so với số lượng các giá trị không thiếu): bạn dùng dictionary để lưu, key là giá trị, value là tỉ lệ %; với cột mà ứng với câu hỏi dạng multichoice, cách làm tương tự như ở trên

Bạn sẽ lưu kết quả vào DataFrame `cate_col_info_df`, trong đó: 
- Tên của các cột là tên của các cột không phải số trong `survey_df`
- Tên của các dòng là: "missing_percentage", "num_values", "value_percentages"  

Để dễ nhìn, tất cả các giá trị bạn đều làm tròn với 1 chữ số thập phân bằng phương thức `.round(1)`.

Gợi ý: có thể bạn sẽ muốn dùng [phương thức `explode`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.Series.explode.html).

```python
# Các cột ứng với câu hỏi khảo sát multichoice
multichoice_cols = ['DevType', 'Gender', 'JobFactors', 
                    'LanguageWorkedWith', 'LanguageDesireNextYear', 
                    'MiscTechWorkedWith', 'MiscTechDesireNextYear', 
                    'NEWCollabToolsWorkedWith', 'NEWCollabToolsDesireNextYear', 
                    'PlatformWorkedWith', 'PlatformDesireNextYear', 
                    'NEWStuck']
```

<!-- grade_id: cell-c1cfa0a88177079d, points: 0, grade: false, solution: true -->
```python
pd.set_option('display.max_colwidth', 100) # Để dễ nhìn
pd.set_option('display.max_columns', None) # Để dễ nhìn

cal_index = ["missing_percentage", "num_values", "value_percentages"]
def get_non_numeric_col(survey_df, multichoice_cols): #Trả về list chứa tên các columns chứa giá trị không phải số và ko multiple
    dtype = survey_df.dtypes
    result = []
    for i in range(len(dtype)):
        if dtype[i] == "object" and dtype.index[i] not in multichoice_cols:
            result.append(dtype.index[i])
    return result

def cate_col_info_df(survey_df, non_numeric_cols, multichoice_cols, cal_index):
    result_df1 = pd.DataFrame(columns = non_numeric_cols, index = cal_index)
    result_df2 = pd.DataFrame(columns = multichoice_cols, index = cal_index)
    for _index, row in result_df1.iterrows():
        for _col in non_numeric_cols:
            if _index == "missing_percentage":
                row[_col] = (survey_df[_col].isna().sum() / survey_df[_col].size * 100).round(1)
            if _index == "num_values":
                series = survey_df[_col]
                num_values = series.value_counts(dropna = True).size
                row[_col] = num_values
            if _index == "value_percentages":
                series = survey_df[_col]
                value_dict = {}
                n = series.value_counts(dropna = True).sum()
                for _i, value in series.value_counts(dropna = True).items():
                    value_dict[_i] = (value / n * 100).round(1)
                sorted_value_dict = dict(sorted(value_dict.items(), key=operator.itemgetter(1),reverse=True))
                row[_col] = sorted_value_dict
    for _index, row in result_df2.iterrows():
        for _col in multichoice_cols:
            if _index == "missing_percentage":
                row[_col] = (survey_df[_col].isna().sum() / survey_df[_col].size * 100).round(1)
            if _index == "num_values":
                series = survey_df[_col]
                series = series.str.split(pat = ";")
                series = series.explode()
                num_values = series.value_counts(dropna = True).size
                row[_col] = num_values
            if _index == "value_percentages":
                series = survey_df[_col]
                series = series.str.split(pat = ";")
                series = series.explode()
                value_dict = {}
                n = series.value_counts(dropna = True).sum()
                for _i, value in series.value_counts(dropna = True).items():
                    value_dict[_i] = (value / n * 100).round(1)
                sorted_value_dict = dict(sorted(value_dict.items(), key=operator.itemgetter(1),reverse=True))
                row[_col] = sorted_value_dict               
    return pd.concat([result_df1, result_df2], axis = 1)

non_numeric_cols = get_non_numeric_col(survey_df, multichoice_cols)
cate_col_info_df = cate_col_info_df(survey_df, non_numeric_cols, multichoice_cols, cal_index)
cate_col_info_df
```

Output:
```
MainBranch  \
missing_percentage                                                                                                  0.5   
num_values                                                                                                            5   
value_percentages   {'I am a developer by profession': 73.5, 'I am a student who is learning to code': 12.5, 'I am n...   

                                     Hobbyist  \
missing_percentage                        0.1   
num_values                                  2   
value_percentages   {'Yes': 78.2, 'No': 21.8}   

                                                                                                                Country  \
missing_percentage                                                                                                  0.6   
num_values                                                                                                          183   
value_percentages   {'United States': 19.5, 'India
```

<!-- grade_id: cell-e748038de60d2ab7, points: 1, grade: true, solution: false -->
```python
# TEST
c = cate_col_info_df['MainBranch']
assert c.loc['missing_percentage'] == 0.5
assert c.loc['num_values'] == 5
assert c.loc['value_percentages']['I am a developer by profession'] == 73.5

c = cate_col_info_df['Hobbyist']
assert c.loc['missing_percentage'] == 0.1
assert c.loc['num_values'] == 2
assert c.loc['value_percentages']['Yes'] == 78.2

c = cate_col_info_df['DevType']
assert c.loc['missing_percentage'] == 23.6
assert c.loc['num_values'] == 23
assert c.loc['value_percentages']['Academic researcher'] == 2.2

c = cate_col_info_df['PlatformWorkedWith']
assert c.loc['missing_percentage'] == 16.5
assert c.loc['num_values'] == 16
assert c.loc['value_percentages']['Docker'] == 10.6
```

---

## Đặt câu hỏi 

Sau khi khám phá dữ liệu, ta đã hiểu hơn về dữ liệu. Bây giờ, ta hãy xem thử có câu hỏi nào có thể được trả lời bằng dữ liệu này.

**Một câu hỏi có thể có là:** Platform nào (Windows, Linux, Docker, AWS, ...) được yêu thích nhất, platform nào được yêu thích nhì, platform nào được yêu thích ba, ...?

Một platform được xem là được yêu thích nếu một người đã dùng platform này (cột "PlatformWorkedWith") và muốn tiếp tục dùng platform trong năm kế (cột "PlatformDesireNextYear").

**Trả lời được câu hỏi này sẽ** phần nào giúp ta định hướng là nên tập trung học platform nào để có thể chuẩn bị cho tương lai (mình nói "phần nào" vì ở đây dữ liệu chỉ giới hạn trong phạm vi những người làm khảo sát của StackOverflow).

---

## Tiền xử lý

Nếu bạn thấy cần thực hiện thêm thao tác tiền xử lý để chuẩn bị dữ liệu cho bước phân tích thì bạn làm ở đây. Bước này là không bắt buộc. 

```python
# YOUR CODE HERE (OPTION) 
```

---

## Phân tích dữ liệu (2.25đ)

Bây giờ, bạn sẽ thực hiện phân tích dữ liệu để trả lời cho câu hỏi ở trên. Cụ thể các bước như sau:
- Bước 1: tính Series `most_loved_platforms`, trong đó: 
    - Index là tên flatform (ở bước khám phá dữ liệu, bạn đã thấy có tất cả 16 platform)
    - Data là tỉ lệ % (từ 0 đến 100, được làm tròn với một chữ số thập phân bằng phương thức `round(1)`) được yêu thích (được sort giảm dần) 
- Bước 2: từ Series `most_loved_platforms`, bạn vẽ bar chart:
    - Bạn cho các bar nằm ngang (cho dễ nhìn)
    - Bạn đặt tên trục hoành là "Tỉ lệ %"

Code bước 1.

<!-- grade_id: cell-6d34c88e2fd46b33, points: 0, grade: false, solution: true -->
```python

platform_index = survey_df['PlatformWorkedWith'].str.split(';').explode().value_counts(dropna = True).index

def replace(col):
    return col.str.split(';')
def most_loved_platforms(survey_df, platform_index):
    df = survey_df[['PlatformWorkedWith', 'PlatformDesireNextYear']].dropna(subset = ['PlatformWorkedWith', 'PlatformDesireNextYear'], how = 'any')
    df.apply(replace)
    df = df.explode('PlatformWorkedWith')
    df = df.explode('PlatformDesireNextYear')
    num_platform = df['PlatformWorkedWith'].value_counts().sum()
    loved = df[df['PlatformWorkedWith'] == df['PlatformDesireNextYear']]
    result = (loved['PlatformWorkedWith'].value_counts() / num_platform * 100)
    return result
most_loved_platforms = most_loved_platforms(survey_df, platform_index)
most_loved_platforms
```

Output:
```
Windows                                            2.651461
Linux                                              1.891842
Linux;Windows                                      1.296910
Android                                            0.759619
Docker;Linux                                       0.568170
                                                     ...   
AWS;Docker;iOS;Linux;MacOS;Raspberry Pi;Windows    0.002059
Heroku;Linux;MacOS;Windows;WordPress               0.002059
AWS;Docker;Linux;MacOS;Raspberry Pi;Windows        0.002059
Kubernetes;MacOS                                   0.002059
Docker;Heroku;Kubernetes                           0.002059
Name: PlatformWorkedWith, Length: 1469, dtype: float64
```

<!-- grade_id: cell-1d22896d0f3f2394, points: 2, grade: true, solution: false -->
```python
# TEST
assert len(most_loved_platforms) == 16
assert most_loved_platforms.loc['Linux'] == 20.2
assert most_loved_platforms.loc['Windows'] == 14.6
assert most_loved_platforms.loc['Docker'] == 12.3
```

Output:
```
AssertionError                            Traceback (most recent call last)
~\AppData\Local\Temp/ipykernel_14092/1951683408.py in <module>
      1 # TEST
----> 2 assert len(most_loved_platforms) == 16
      3 assert most_loved_platforms.loc['Linux'] == 20.2
      4 assert most_loved_platforms.loc['Windows'] == 14.6
      5 assert most_loved_platforms.loc['Docker'] == 12.3

AssertionError:
```

Code bước 2.

<!-- grade_id: cell-5395c611d6065339, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Bạn đã hiểu tại sao mình khuyên bạn là nên tập làm quen dần với các câu lệnh của Linux chưa 😉 

---

## Đặt câu hỏi của bạn (1.5đ)

Bây giờ, đến lượt bạn phải tự suy nghĩ và đưa ra câu hỏi mà có thể trả lời bằng dữ liệu. Ngoài việc đưa ra câu hỏi, bạn cũng phải giải thích để người đọc thấy nếu trả lời được câu hỏi thì sẽ có lợi ích gì. Bạn nên sáng tạo một xíu, không nên đưa ra câu hỏi cùng dạng với câu hỏi của mình ở trên.  

<!-- grade_id: cell-2e7a92ae99ca7449, points: 1.5, grade: true, solution: true -->
YOUR ANSWER HERE

---

## Tiền xử lý để chuẩn bị dữ liệu cho bước phân tích để trả lời cho câu hỏi của bạn

Phần này là không bắt buộc. 

```python
# YOUR CODE HERE (OPTION) 
```

---

## Phân tích dữ liệu để ra câu trả lời cho câu hỏi của bạn (2đ)

<!-- grade_id: cell-83e11892184cf811, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```